# SLU 09 - Exercise notebook 

Before you start:
>- This notebook is divided into 8 questions: 
    - Question 1 to 6 are easy, practice questions, expect that part to take about 2 hours. 
    - Questions 7 and 8 are hard question, expect to spend at least another 3 hours on them. They are close to the level of difficulty to pass the Academy test. 
    
> - If you are unsure about your answer, run the asserts (the tests). Even if it is a bit incomplete, run the asserts, and see if they give you clues. 
    - In the last 2 questions the asserts will be a bit less helpful, so that you can get some practice of dealing with genuine python errors. 
    - Read the instructions carefully. Everything that is there is critical
    - Don't worry if these concepts feel weird at first. Classes are not easy, you are now getting into advanced stuff! 
    
> - If you are stuck, remember to:
    - Play around! Open a new cell, and mess around with the objects. Instantiate them and see if they complain. See what they print, what they return, etc. Remember, in real life there are no asserts to guide you!  
    - Always read the error messages carefully. They should be good clues about what is wrong. 
    - As always, don't be afraid to ask questions! 


----

### Exercise 1 - Playing around with objects 

This exercise is very simple. We are just going to make a couple new fruits, like we did in the first Learning Notebook. 

Make the following two fruits: 
- pears, which is called "pears", has 4 units, costs 3 euros per unit, and expires in 14 days
- tangerines, which is called "tangerines", has 2 units, costs 6 euros per unit, and expires in 2 days

This exercise is very easy, just look up how it was done in the Learning notebook and do a couple of them here. 

Reminder: the Fruit class has the following attributes: 
- `price_per_unit`
- `name`
- `nr_units`
- `days_until_expired`

> Protip: Note that you can use the tab auto-complete to help you.  
> Try writing `Fruit(pr` and hitting tab. It should auto-complete to price! 


<!-- grade_id: cell-4640dc55a4c658cd, points: 0, grade: false, solution: true -->
```python
from utils import Fruit

# pears = ...
# tangerines = ...

# YOUR CODE HERE
pears = Fruit(price_per_unit=3,
             name="pears",
             nr_units=4,
             days_until_expired=14)
tangerines =Fruit(name="tangerines",
                 nr_units=2,
                 price_per_unit=6,
                 days_until_expired=2)
```

<!-- grade_id: cell-f4f479ecb4b7c3ef, points: 1, grade: true, solution: false -->
```python
from utils import exercise_1_grading
assert (pears and tangerines), 'Did you forget to make one of the fruits?'
exercise_1_grading(pears, tangerines)
print('Great success! Exercise passed!')
```

Output:
```
Great success! Exercise passed!
```

----
### Exercise 2 - add up the price of fruits 

Great! For your second exercise you will add the price of all the fruits: 
- apples
- bananas 
- oranges 
- pears
- tangerines 

The price is the product of the price per unit and number of units. But remember, the Fruit object already comes with a `calculate_price` method! 

Create a function called `calculate_price_of_all_fruits`, which takes a list called `all_fruits` that contains fruit objects and returns the total price. We've already filled in some things for you. 

Remember: the function should take any list of fruits, don't hard code any particular list! 

<!-- grade_id: cell-d5a600cefd3c41c7, points: 0, grade: false, solution: true -->
```python
def calculate_price_of_all_fruits(all_fruits): 
    """
    Calculates the price for every fruit in the all_fruits list, 
        and returns the total price of fruits in the list
    Args: 
        all_fruits (list): a list with fruit objects 
        
    Returns: 
        total_price (int): the total price of the fruits in the list
    """
    total_price = 0
    for i in all_fruits:
        total_price += i.calculate_price()
    return total_price
        
    # hint: this may require a for loop, and a variable to keep track of the total price 

    # YOUR CODE HERE

    
# this should return 24 
print(calculate_price_of_all_fruits([pears, tangerines]))
```

Output:
```
24
```

<!-- grade_id: cell-ee1a997f801ac0bb, points: 1, grade: true, solution: false -->
```python
from utils import exercise_2_grading
exercise_2_grading(calculate_price_of_all_fruits)
print('Great success! Exercise passed!')
```

Output:
```
Great success! Exercise passed!
```

---
### Exercise 3 - Using objects with other objects 

In this exercise you will do something very similar to the learning notebook. You will simply instantiate a basket, create some toilet paper, and then will add and remove stuff until you have the following: 

- Total value of the basket should be between 30 and 35 
- the basket must have at least 3 types of fruit
- the basket must not contain any items that will expire in the next 8 days 
- the basket must have at least some toilet paper. It should be ... 
    - named "Soft but cheap" 
    - cost 4 euros
    - have "Double leaf". 
    - The expiration date and number of units are up to you, they won't be tested

Reminders: 
```
- You already have some instance variables to work with: 
    - apples 
    - oranges
    - bananas 
    - pears 
    - tangerines 
- You only have the Class variable Toiletpaper to work with. 
- You'll need to create some toilet paper before you can use it. 
    - Call the variable toilet_paper.
    - Give it nr_units = 1 
```

_Note: obviously you can figure out exactly what you need to add ahead of time without trial and error. However, it's good practice to just use the objects. Try adding, checking the price, removing, etc until you have an answer that meets all the conditions._

<!-- grade_id: cell-cffd2f4f2528de89, points: 0, grade: false, solution: true -->
```python
from utils import Toiletpaper, Basket, get_fruits
apples, bananas, oranges = get_fruits()  # you already have some fruits

# my_basket = ... 
# toilet_paper = ...
# my_basket.<do something> 

# YOUR CODE HERE

my_basket = Basket() 
toilet_paper = Toiletpaper(name="Soft but cheap",
                          nr_units=1,
                          price_per_unit=4,thickness="Double leaf",
                          days_until_expired=1000000000)

my_basket.add_item(apples) 
my_basket.add_item(oranges) 
my_basket.add_item(bananas) 
my_basket.add_item(toilet_paper)

my_basket.examine_basket()
my_basket.check_for_items_close_to_expire(8)

my_basket.remove_item(bananas)
my_basket.add_item(pears)

my_basket.examine_basket()
```

Output:
```
The total price is 32
- 10 apples (total price 10)
- 2 oranges (total price 6)
- 6 bananas (total price 12)
- 1 Soft but cheap (total price 4)
The item bananas will expire in 7 days
The total price is 32
- 10 apples (total price 10)
- 2 oranges (total price 6)
- 1 Soft but cheap (total price 4)
- 4 pears (total price 12)
```

<!-- grade_id: cell-277ef4e6f1a47d0d, points: 1, grade: true, solution: false -->
```python
from utils import exercise_3_grading 
exercise_3_grading(my_basket, toilet_paper)
print('Great success! Exercise passed!')
```

Output:
```
Great success! Exercise passed!
```

----
### Exercise 4: Making a basic object 

Ok, here is what is going ot happen. We're gonna move house. And we're going to have to pack all our stuff. 

We will take a bunch of `Item`s in our house, and put them into `Boxes`. 

Let's start with creating some `Item`s. Here is what you need to do: 

Create an `Item` class, which will have 3 attributes: 
- name 
- weight 
- fragile (this will be boolean)
- broken (this will be boolean, always starts out as False) 

Create the following items: 
- `chair` which will weigh 15, be called "kitchen chair", and not be fragile 
- `vase` which will weigh 3, be called "ming vase", and be fragile 
- `pillow` which will weigh 1, be called "bed pillow", and not be fragile  

<!-- grade_id: cell-2e5ef92f86968bc5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Item:
    def __init__(self,name,weight,fragile:bool):
        self.name = name
        self.weight = weight
        self.fragile = fragile
        self.broken = False
        
chair = Item('kitchen chair', weight=15,fragile=False)
vase = Item('ming vase', weight=3,fragile=True)
pillow = Item('bed pillow', 1,False)
```

<!-- grade_id: cell-61f3abe376b0ed6d, points: 1, grade: true, solution: false -->
```python
from utils import exercise_4_grading 
assert Item, "please create the Item class!"
assert all([vase, chair, pillow])
exercise_4_grading(Item, chair, vase, pillow)
print('Great success! Section passed!')
```

Output:
```
Great success! Section passed!
```

---
### Exercise 5: working with methods 

For this exercise we will create a `BasicBox` class (note: we'll create a `Box` class in the next exercise). This box will allow you to add and remove `Item`s, and will have a weight limit.

First step, let's create the `BasicBox` object. It will have the following attributes: 
- `contents`: a list that will start empty. We will use it to keep track of items. 
- `weight_limit`: a number specific to each new box we create (different `Box`es will have different weight limits)
- `name`: a string name 
- `broken`: boolean, whether the box has been broken by having too much weight. Starts as False. 
- `weight`: defaults to zero (the weight of the box is negligible) 

The `BasicBox` should also have a method called `add_item` that takes an `item` and adds it to the `contents`. 

<!-- grade_id: cell-a9722b9debaeb161, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class BasicBox:
    def __init__(self,weight_limit,name):
        self.contents = []
        self.weight_limit = weight_limit
        self.name = name
        self.broken = False
        self.weight = 0
        
    def add_item(self,item):
        self.contents.append(item)
        
        

```

<!-- grade_id: cell-a9853f7b4e57fe87, points: 2, grade: true, solution: false -->
```python
from utils import exercise_5_grading 
assert BasicBox
exercise_5_grading(BasicBox)
print('Great success! Exercise passed!')
```

Output:
```
Great success! Exercise passed!
```

---
### Exercise 6: calling methods with other methods 
Ok, our box is a bit too basic. Let's make a proper box (feel free to copy your code from the previous exercise as this one is an extension of it) 

This time we'll create a `Box` object, that allows us to add items, but has a few more methods: 

- `calculate_total_weight`: calculates the weight of all `items` in the `contents`, and updates the `weight` attribute. 

- `check_if_weight_limit_exceeded`: this method will call `calculate_total_weight`, then check if the total weight is over the box's `weight_limit`. If it has been exceeded, the box should change its `broken` attribute to `True`, and `break_all_fragile_items`

- `break_all_fragile_items`: change the `broken` attribute of all `fragile` items in the `contents` to `True` (don't break the non-fragile items!) 

_Note 1: don't return the total_weight, update the `self.weight` attribute of the box itself._  
_Note 2: you will probably want to use some of your Items from exercise 4 to test whether your Box does what is expected._

```python
            
    def check_if_weight_limit_exceeded(self):
        if self.calculate_total_weight() > self.weight_limit :
            print("too heavy")
    
    def break_all_fragile_items(self):
        for item in self.contents:
            if item.fragile == True :
                item.broken = True
            
```

<!-- grade_id: cell-ef56f4c9a022e15a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Box:
    def __init__(self,weight_limit,name):
        self.contents = []
        self.weight_limit = weight_limit
        self.name = name
        self.broken = False
        self.weight = 0
        
    def add_item(self,item):
        self.contents.append(item)
        
    def calculate_total_weight(self):
        self.weight = 0
        for item in self.contents:
            self.weight += item.weight
        
    
    def break_all_fragile_items(self):
        for item in self.contents:
            if item.fragile:
                self.broken = True
                item.broken = True
                
    def check_if_weight_limit_exceeded(self):
        self.calculate_total_weight()
        if self.weight > self.weight_limit :
            self.break_all_fragile_items()
            
    def describe_every_item(self):
        for item in self.contents:
            # print some info about the item 
            print('- {0}. Is it fragile? {1} (broken: {2})'.format(
                item.name, str(item.fragile), str(item.broken)))

            
            
box1 = Box(name= 'box1',weight_limit=15)

for item in [chair,vase,pillow]:
    box1.add_item(item)

box1.calculate_total_weight()
box1.check_if_weight_limit_exceeded()
box1.describe_every_item()
```

Output:
```
- kitchen chair. Is it fragile? False (broken: False)
- ming vase. Is it fragile? True (broken: True)
- bed pillow. Is it fragile? False (broken: False)
```

<!-- grade_id: cell-9eae48817b9d3b7f, points: 4, grade: true, solution: false -->
```python
from utils import exercise_6_grading
assert Box, "Implement Box!"
exercise_6_grading(Box)
print('Great success! Section passed!')
```

Output:
```
Great success! Section passed!
```

----
> Before you continue:
- _The exercises up to here were good for practice. The last two are hard ones._ 
- The last 2 exercises have no new material. They will just have less hand-holding and be more demanding. If you get in trouble, see your own solutions above for clues!
- To give you a more realistic idea of debugging in the real world we won't tell you where your mistakes are, but simply whether you are passing the tests (whether it does what is expected). You will have to use the errors that python provides, as will be the case in the real world. 
- You can of course ask for help at any time, but try suffering through these two on your own if you can, it's very good practice for the real world :)  
- _Expect to take quite a while (probably over 2 hours) to get exercise 7 and 8 right._
- _It's often useful to draft the "skeleton" of your code before jumping in. What are the steps? What are the methods of each class you will need? Sometimes writing on a piece of paper is a good support._ 
- _This section will test you on what you learned, but also includes functions, and other things you have learned before. You may need to check a few past SLUs for this one_

- _After you create a class, make an instance, see if it behaves the way you expect. It's crucial to gain the habit of testing your stuff as you go, to see if it is working._
    - it's your responsibility to create instances, find the actual error, and try to fix it!_ 
- _This is the last block in this SLU! If you got this, you are getting pretty good at objects, and are probably ready to deal with the Academy entry exercise that has this SLU's material._

---
### Exercise 7: let's make a SportsBall league! 

You know those games where there are professional sportsball leagues, with players, teams and competitions? Yeah, neither do I, and it will show in this exercise. 

--------

We are going to create two classes. 


**Player class**
`Player` class, which will have the following attributes: 
- `name`, which is passed as an argument 
- `salary`, which is passed as an argument
- `accuracy`, which is passed as an argument

It will also have a method called `attempt_to_score`, which takes no arguments. This method will return `True` if the player's accuracy is higher than a random number generated between 0 and 10 (you will generate it using the code explained next), otherwise it will return `False`. 

To generate a random number between 0 and 10 use the following: `np.random.randint(0, 10)`

**Team class**
`Team` class, which will have the following attributes: 
- `name`, which is passed as an argument 
- `budget`, which is passed as an argument 
- `players`, which defaults to an empty list 
- `wins`, which defaults to 0
- `games_played`, which defaults to 0
- `points_in_game`, which defaults to 0

It will have the following attributes: 
- `hire_player`, takes a player as attribute. This method should check if the salary is equal to or lower than the budget, and if so add the player to the list of players, and remove their salary from the budget. If the team does not have enough budget to hire the player you should print something that helps you know that happened (your choice). 
- `fire_player`, takes a player as attribute. This method should remove the player from the list of players, and add their salary back to the budget. 
- `reset_points`, takes no arguments. This method should set the `points_in_game` back to 0 
- `all_players_attempt_to_score`, takes no arguments. This method should start by calling the `reset_points` method, then loop over every player on the team and have each of them attempt to score (using the player's `attempt_to_score` method). For each one that scores it should increment `points_in_game` by 1. 

<!-- grade_id: cell-05a39eba6e7e96c1, points: 0, grade: false, solution: true -->
```python
import numpy as np  # we will need this for attempts to score 

# YOUR CODE HERE
class Player:
    def __init__(self, name, salary, accuracy):
        self.name = name
        self.salary = salary
        self.accuracy = accuracy
        
    def attempt_to_score(self):
        random_number = np.random.randint(0, 10)
        if self.accuracy > random_number :
            return True
        else:
            return False
        




class Team:
    def __init__(self, name, budget):
        self.name = name
        self.budget = budget
        self.players = []
        self.wins = 0
        self.games_played = 0
        self.points_in_game = 0
        
    def hire_player(self,player):
        if player.salary <= self.budget :
            self.players.append(player)
            self.budget -= player.salary
            print("Player {0} was hired".format(player.name))
        else:
            print("Player {0} is too expensive and cannot be hired".format(player.name))
        
    def team_players(self):
        print('This are the team players:')
        for players in self.players:
            # print some info about the item 
            print('- {0}'.format(
                players.name))
            
    def fire_player(self, player):
        self.players.remove(player)
        self.budget += player.salary
        
    def reset_points(self):
        self.points_in_game = 0
    
    def all_players_attempt_to_score(self):
        self.reset_points()
        for player in self.players:
            if player.attempt_to_score() == True:
                self.points_in_game += 1
                    
    def team_score(self):
        print('This is the team score:{0}'.format(self.points_in_game))        
 
 
Johnny = Player(name= "Johnny", salary=100, accuracy= 0)
a = Player(name= "a", salary=102, accuracy= 0)
b = Player(name= "b", salary=3, accuracy= 0)
c = Player(name= "c", salary=1000, accuracy= 0)
d = Player(name= "d", salary=100, accuracy= 0)


Manchester = Team(name = "Manchester", budget = 10000)

for i in [Johnny,a,b,c,d]:
    Manchester.hire_player(i)


print(Manchester.name)
print(Manchester.budget)


print(Manchester.budget)
Manchester.team_players()


Manchester.all_players_attempt_to_score()

Manchester.team_score()


```

Output:
```
Player Johnny was hired
Player a was hired
Player b was hired
Player c was hired
Player d was hired
Manchester
8695
8695
This are the team players:
- Johnny
- a
- b
- c
- d
This is the team score:0
```

```python
len(Manchester.players)
```

Output:
```
5
```

<!-- grade_id: cell-e429403c7f5eff78, points: 5, grade: true, solution: false -->
```python
from utils import exercise_7_validate_Player, exercise_7_validate_Team
import numpy as np 
assert Player, team
exercise_7_validate_Player(Player)
exercise_7_validate_Team(Team, Player)
print('Great success! Section passed!')
```

Output:
```
Player test was hired
Player test 1 was hired
Player test 2 was hired
Player test is too expensive and cannot be hired
Player test 1 was hired
Player test 2 is too expensive and cannot be hired
Player test 1 was hired
Player test 2 was hired
Player test1 was hired
Player test1 was hired
Player test1 was hired
Great success! Section passed!
```

---
### Exercise 8: making a League for your Teams to compete in  

Almost there! For this last exercise we will create a League. 

The `League` class will have the following attributes:
- `name`, which is passed as an argument 
- `teams`, which defaults to an empty list 
- `games`, which defaults to an empty list 

It will have the following methods: 
- `add_team`, which takes a team as an argument. This method will check if the team has exactly 5 players, and if so will add it to the `teams` list. If it does not have exactly 5 players the team will not get added, and you should print something to tell you that happened. 

- `declare_game_winner`, which takes two arguments: `team_1` and `team_2`. This method will compare the number of `points_in_game` between the two teams, and will give 1 `win` to the team with the most `points_in_game`. If they have the same number of points, neither gets a `win`. 

- `run_game`, which takes two arguments: `team_1` and `team_2`. This method will make the two teams play a game, which means they will both have `all_players_attempt_to_score`, and then you will increase their `games_played` by 1. After that has happened, you will `declare_game_winner` between the two. 

- `print_current_scores`, takes no arguments. This method will loop over the teams, and will print how many points they have in an easy to read way (you can chose how, the exact wording isn't tested). 

_Note: if you can make this `League` work with your `Team`s and your `Player`s you will have made your first proper bit of code with multiple classes, and are ready for the next SLU. Expect this to take a bit of time, but we promise it's worth suffering through :)_

```python
Mohnny = Player(name= "Mohnny", salary=100, accuracy= 10)
a = Player(name= "a", salary=102, accuracy= 10)
b = Player(name= "b", salary=3, accuracy= 10)
c = Player(name= "c", salary=1000, accuracy= 10)
d = Player(name= "d", salary=100, accuracy= 10)


Chelsea = Team(name = "Chelsea", budget = 10000)

for i in [Mohnny,a,b,c,d]:
    Chelsea.hire_player(i)
```

Output:
```
Player Mohnny was hired
Player a was hired
Player b was hired
Player c was hired
Player d was hired
```

<!-- grade_id: cell-a43a4e57177eca9d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class League:
    def __init__(self, name):
        self.name = name
        self.teams = []
        self.games = []
        
    def add_team(self,team):
        if len(team.players) == 5:
            self.teams.append(team)
            print("Team {0} is now on the {1} league".format(team.name,self.name))
        else:
            print("Team {0} has not 5 players".format(team.name))
    
    def declare_game_winner(self, team_1,team_2):
        if team_1.points_in_game > team_2.points_in_game:
            team_1.wins += 1
            print("Team {0} has {1} win".format(team_1.name,team_1.wins))
        elif team_1.points_in_game < team_2.points_in_game:
            team_2.wins += 1
            print("Team {0} has {1} win".format(team_2.name, team_2.wins))
            
    def run_game(self, team_1,team_2):
        team_1.all_players_attempt_to_score()
        team_1.games_played +=1
        print("Team {0} has {1} points".format(team_1.name,team_1.points_in_game))
        team_2.all_players_attempt_to_score()
        team_2.games_played +=1
        print("Team {0} has {1} points".format(team_2.name,team_2.points_in_game))
        self.declare_game_winner(team_1,team_2)
    
    def print_current_scores(self):
        for team in self.teams:
            print("Team {0} has {1} points in the {2} League".format(team.name, team.wins, self.name) )
        
        
        
            
    
        
            
            

            
Champions = League(name="Champions")

for i in [Chelsea, Manchester]:
    Champions.add_team(team = i)



            
```

Output:
```
Team Chelsea is now on the Champions league
Team Manchester is now on the Champions league
```

```python

Champions.run_game(team_1 = Manchester,team_2 = Chelsea)
Champions.declare_game_winner(team_1 = Manchester, team_2 = Chelsea)

Champions.print_current_scores()
```

Output:
```
Team Manchester has 0 points
Team Chelsea has 5 points
Team Chelsea has 1 win
Team Chelsea has 2 win
Team Chelsea has 2 points in the Champions League
Team Manchester has 0 points in the Champions League
```

<!-- grade_id: cell-1d7bb58d538228da, points: 5, grade: true, solution: false -->
```python
from utils import exercise_8_validate_League
exercise_8_validate_League(League, Team, Player)
print("Great Success! Section passed!")
```

Output:
```
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Team test_team is now on the test league
Team test_team is now on the test league
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Team test_team is now on the test league
Team test_team has not 5 players
Team test_team has not 5 players
Team test_team has 1 win
Team test_team has 2 win
Team test_team has 1 win
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Team test_team is now on the test league
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Player test was hired
Team test_team is now on the test l
```

# Submit your work!

To submit your work, [follow the steps here, in the step "Grading the Exercise Notebook"!](https://github.com/LDSSA/ds-prep-course-2023#22---working-on-the-learning-units)

```python

```

```python

```