# SLU 17_1 - Exam Prep I

## Batch 1 Python Exam 

This is the Python admission exam from the Lisbon Data Science Starters Academy - Batch 1. The allocated time for this exam was 2 hours. If you plan to take the admission exam this year, it's a good idea to measure the time you need to solve this notebook. You can see how far you get in 2 hours and also how much time you need to complete the notebook. At the same time, it's a good idea to not rush through the exercises and rather concentrate on working thoroughly.
At the end, you submit the notebook to the portal as usual. You can submit as many times as you like.

```python
#used for evaluation
import inspect
import math
import re
import json
import sys
import hashlib

def _hash(s):
    return hashlib.blake2b(
        bytes(str(s), encoding='utf8'),
        digest_size=5
    ).hexdigest()
```

## Question 1

Complete the function `import_data`. This function should:
- Open the **countries.json** file that is located within the folder name **data** ;
- Use Python’s module json to load the data from the file into a dictionary;
- Return the dictionary;

```python
!ls data
```

Output:
```
countries.json
```

```python
with open("./data/countries.json") as user_file:
    parsed_json = json.load(user_file)

print(parsed_json)
```

Output:
```
{'Afghanistan': {'Continent': 'Asia', 'Population': 31056997.0, 'Area': 647500, 'GDP': 700.0, 'Literacy': 36.0, 'Phones': 3.22}, 'Albania': {'Continent': 'Europe', 'Population': 3581655.0, 'Area': 28748, 'GDP': 4500.0, 'Literacy': 86.5, 'Phones': 71.2}, 'Algeria': {'Continent': 'Africa', 'Population': 32930091.0, 'Area': 2381740, 'GDP': 6000.0, 'Literacy': 70.0, 'Phones': 78.1}, 'American Samoa': {'Continent': 'Oceania', 'Population': 57794.0, 'Area': 199, 'GDP': 8000.0, 'Literacy': 97.0, 'Phones': 259.54}, 'Andorra': {'Continent': 'Europe', 'Population': 71201.0, 'Area': 468, 'GDP': 19000.0, 'Literacy': 100.0, 'Phones': 497.18}, 'Angola': {'Continent': 'Africa', 'Population': 12127071.0, 'Area': 1246700, 'GDP': 1900.0, 'Literacy': 42.0, 'Phones': 7.78}, 'Anguilla': {'Continent': 'America', 'Population': 13477.0, 'Area': 102, 'GDP': 8600.0, 'Literacy': 95.0, 'Phones': 460.04}, 'Antigua & Barbuda': {'Continent': 'America', 'Population': 69108.0, 'Area': 443, 'GDP': 11000.0, 'Literacy': 
```

<!-- grade_id: cell-ce9179c0612fd6b0, points: 0, grade: false, solution: true -->
```python
def import_data():
    with open("./data/countries.json") as user_file:
        parsed_json = json.load(user_file)
    return parsed_json
    
    
imp = import_data()

print(imp)
```

Output:
```
{'Afghanistan': {'Continent': 'Asia', 'Population': 31056997.0, 'Area': 647500, 'GDP': 700.0, 'Literacy': 36.0, 'Phones': 3.22}, 'Albania': {'Continent': 'Europe', 'Population': 3581655.0, 'Area': 28748, 'GDP': 4500.0, 'Literacy': 86.5, 'Phones': 71.2}, 'Algeria': {'Continent': 'Africa', 'Population': 32930091.0, 'Area': 2381740, 'GDP': 6000.0, 'Literacy': 70.0, 'Phones': 78.1}, 'American Samoa': {'Continent': 'Oceania', 'Population': 57794.0, 'Area': 199, 'GDP': 8000.0, 'Literacy': 97.0, 'Phones': 259.54}, 'Andorra': {'Continent': 'Europe', 'Population': 71201.0, 'Area': 468, 'GDP': 19000.0, 'Literacy': 100.0, 'Phones': 497.18}, 'Angola': {'Continent': 'Africa', 'Population': 12127071.0, 'Area': 1246700, 'GDP': 1900.0, 'Literacy': 42.0, 'Phones': 7.78}, 'Anguilla': {'Continent': 'America', 'Population': 13477.0, 'Area': 102, 'GDP': 8600.0, 'Literacy': 95.0, 'Phones': 460.04}, 'Antigua & Barbuda': {'Continent': 'America', 'Population': 69108.0, 'Area': 443, 'GDP': 11000.0, 'Literacy': 
```

```python

type(imp)
```

Output:
```
dict
```

<!-- grade_id: cell-78d54c806fb28266, points: 2, grade: true, solution: false -->
```python
countries = import_data()
assert isinstance(countries, dict), "Not correct. Keep trying."
assert len(countries) == 194, "Not correct. Keep trying."
assert countries["Albania"]["Continent"] == "Europe", "Not correct. Keep trying."
assert countries["Vietnam"]["Area"] == 329560, "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries_backup = import_data()
```

Output:
```
---- All asserts passed ----
```

## Question 2

The dictionary that you built contains information about a serie of countries. So from now on, we’ll call it the `countries` dictionary.
In the `countries` dictionary, each country has a field called `Population` (be aware that this is case sensitive) that gives the number of people living in that country. Since this is a count, the data should have integer type. However, we have found some numbers with type of *float*. Let's try to fix this!

Complete the function `convert_population_to_int`. This function should:
- Receive the `countries` dictionary as argument;
- Convert the `Population` field from float to int type for all the countries, using a `for` loop;
- Return the updated `countries` dictionary;

```python
print(countries)
```

Output:
```
{'Afghanistan': {'Continent': 'Asia', 'Population': 31056997.0, 'Area': 647500, 'GDP': 700.0, 'Literacy': 36.0, 'Phones': 3.22}, 'Albania': {'Continent': 'Europe', 'Population': 3581655.0, 'Area': 28748, 'GDP': 4500.0, 'Literacy': 86.5, 'Phones': 71.2}, 'Algeria': {'Continent': 'Africa', 'Population': 32930091.0, 'Area': 2381740, 'GDP': 6000.0, 'Literacy': 70.0, 'Phones': 78.1}, 'American Samoa': {'Continent': 'Oceania', 'Population': 57794.0, 'Area': 199, 'GDP': 8000.0, 'Literacy': 97.0, 'Phones': 259.54}, 'Andorra': {'Continent': 'Europe', 'Population': 71201.0, 'Area': 468, 'GDP': 19000.0, 'Literacy': 100.0, 'Phones': 497.18}, 'Angola': {'Continent': 'Africa', 'Population': 12127071.0, 'Area': 1246700, 'GDP': 1900.0, 'Literacy': 42.0, 'Phones': 7.78}, 'Anguilla': {'Continent': 'America', 'Population': 13477.0, 'Area': 102, 'GDP': 8600.0, 'Literacy': 95.0, 'Phones': 460.04}, 'Antigua & Barbuda': {'Continent': 'America', 'Population': 69108.0, 'Area': 443, 'GDP': 11000.0, 'Literacy': 
```

```python
'''for country, info in countries.items():
    for key in info:
        info['Population'] = int(info['Population'])
        
print(countries'''
```

Output:
```
"for country, info in countries.items():\n    for key in info:\n        info['Population'] = int(info['Population'])\n        \nprint(countries"
```

<!-- grade_id: cell-28d89deeef261406, points: 0, grade: false, solution: true -->
```python
def convert_population_to_int(countries):
    # YOUR CODE HERE
    for country, info in countries.items():
        for key in info:
            info['Population'] = int(info['Population'])
    return countries
```

```python
x = convert_population_to_int(countries)
```

```python
print(x)
```

Output:
```
{'Afghanistan': {'Continent': 'Asia', 'Population': 31056997, 'Area': 647500, 'GDP': 700.0, 'Literacy': 36.0, 'Phones': 3.22}, 'Albania': {'Continent': 'Europe', 'Population': 3581655, 'Area': 28748, 'GDP': 4500.0, 'Literacy': 86.5, 'Phones': 71.2}, 'Algeria': {'Continent': 'Africa', 'Population': 32930091, 'Area': 2381740, 'GDP': 6000.0, 'Literacy': 70.0, 'Phones': 78.1}, 'American Samoa': {'Continent': 'Oceania', 'Population': 57794, 'Area': 199, 'GDP': 8000.0, 'Literacy': 97.0, 'Phones': 259.54}, 'Andorra': {'Continent': 'Europe', 'Population': 71201, 'Area': 468, 'GDP': 19000.0, 'Literacy': 100.0, 'Phones': 497.18}, 'Angola': {'Continent': 'Africa', 'Population': 12127071, 'Area': 1246700, 'GDP': 1900.0, 'Literacy': 42.0, 'Phones': 7.78}, 'Anguilla': {'Continent': 'America', 'Population': 13477, 'Area': 102, 'GDP': 8600.0, 'Literacy': 95.0, 'Phones': 460.04}, 'Antigua & Barbuda': {'Continent': 'America', 'Population': 69108, 'Area': 443, 'GDP': 11000.0, 'Literacy': 89.0, 'Phones': 
```

<!-- grade_id: cell-33b501c79f675ab0, points: 2, grade: true, solution: false -->
```python
countries = convert_population_to_int(countries)
assert len(countries) == 194, "Not correct. Keep trying."
source = inspect.getsource(convert_population_to_int)
assert "for" in source, "Not correct. Keep trying."
assert _hash(str(countries)) == "4e2563f26e", "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries = countries_backup
```

Output:
```
---- All asserts passed ----
```

## Question 3

In the countries dictionary, each country has a field called `Area`. This is the area of the country given in square miles.

Complete the function `convert_area_to_sq_km`. This function should:
- Receive the `countries` dictionary as argument;
- Convert the `Area` field from square miles to square kilometres;
 - Use the following conversion: 1 sq mi = 2.58999 sq km;
 - Round the result to 1 decimal digit;
- Return the updated `countries` dictionary;

```python
'''for country, info in countries.items():
    for key in info:
        #info['Area'] = float(info['Area'])
        print(country,':',key, ':', info[key]'''
```

Output:
```
"for country, info in countries.items():\n    for key in info:\n        #info['Area'] = float(info['Area'])\n        print(country,':',key, ':', info[key]"
```

```python
'''for country, info in countries.items():
    info['Area'] = round((info['Area'] * 2.58999),1'''
   
        
        
```

Output:
```
"for country, info in countries.items():\n    info['Area'] = round((info['Area'] * 2.58999),1"
```

```python

```

```python
print(countries)
```

Output:
```
{'Afghanistan': {'Continent': 'Asia', 'Population': 31056997.0, 'Area': 647500, 'GDP': 700.0, 'Literacy': 36.0, 'Phones': 3.22}, 'Albania': {'Continent': 'Europe', 'Population': 3581655.0, 'Area': 28748, 'GDP': 4500.0, 'Literacy': 86.5, 'Phones': 71.2}, 'Algeria': {'Continent': 'Africa', 'Population': 32930091.0, 'Area': 2381740, 'GDP': 6000.0, 'Literacy': 70.0, 'Phones': 78.1}, 'American Samoa': {'Continent': 'Oceania', 'Population': 57794.0, 'Area': 199, 'GDP': 8000.0, 'Literacy': 97.0, 'Phones': 259.54}, 'Andorra': {'Continent': 'Europe', 'Population': 71201.0, 'Area': 468, 'GDP': 19000.0, 'Literacy': 100.0, 'Phones': 497.18}, 'Angola': {'Continent': 'Africa', 'Population': 12127071.0, 'Area': 1246700, 'GDP': 1900.0, 'Literacy': 42.0, 'Phones': 7.78}, 'Anguilla': {'Continent': 'America', 'Population': 13477.0, 'Area': 102, 'GDP': 8600.0, 'Literacy': 95.0, 'Phones': 460.04}, 'Antigua & Barbuda': {'Continent': 'America', 'Population': 69108.0, 'Area': 443, 'GDP': 11000.0, 'Literacy': 
```

<!-- grade_id: cell-3fd27539b41bed5c, points: 0, grade: false, solution: true -->
```python
def convert_area_to_sq_km(countries):
    # YOUR CODE HERE
    for country, info in countries.items():
        info['Area'] = round((info['Area'] * 2.58999),1)
    return countries
```

<!-- grade_id: cell-b77708fde98d925a, points: 2, grade: true, solution: false -->
```python
countries = convert_area_to_sq_km(countries)
for country in countries.keys():
    assert isinstance(countries[country]["Area"], float), \
            "Not correct. Keep trying."
assert len(countries) == 194, "Not correct. Keep trying."
assert math.isclose(countries["Haiti"]["Area"], 71872.2, abs_tol=0.1), \
       "Not correct. Keep trying."
assert math.isclose(countries["Morocco"]["Area"], 1156560.0, abs_tol=0.1), \
       "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries = countries_backup
```

Output:
```
---- All asserts passed ----
```

## Question 4

Complete the function `get_europe_countries`. This function should:
- Receive the `countries` dictionary as argument;
- Build a list with the names of the countries in Europe sorted by alphabetical order (from A to Z);
 - Use the field called `Continent`;
 - Use list comprehension;
 - Use the sort method from the list data type;
- Return the list;

```python
'''list_europe = []
for country, info in countries.items():
    if info['Continent'] == 'Europe':
        list_europe.append(country'''
```

Output:
```
"list_europe = []\nfor country, info in countries.items():\n    if info['Continent'] == 'Europe':\n        list_europe.append(country"
```

```python
'''list_europ'''
```

Output:
```
'list_europ'
```

```python

```

```python
'''list_europe = []
list_europe = [list_europe.append(country) for country, info in countries.items() if info['Continent'] == 'Europe']
   '''
```

Output:
```
"list_europe = []\nlist_europe = [list_europe.append(country) for country, info in countries.items() if info['Continent'] == 'Europe']\n   "
```

```python
'''list_europ'''
```

Output:
```
'list_europ'
```

<!-- grade_id: cell-ee1f38ef730c2004, points: 0, grade: false, solution: true -->
```python
def get_europe_countries(countries):
    list_europe = []
    list_europe = [country for country, info in countries.items() if info['Continent'] == 'Europe']
    list_europe.sort()
    return list_europe
```

```python
contries_europe = get_europe_countries(countries)


```

```python
print(contries_europe)
```

Output:
```
['Albania', 'Andorra', 'Austria', 'Belgium', 'Bulgaria', 'Croatia', 'Czech Republic', 'Denmark', 'Estonia', 'Finland', 'France', 'Germany', 'Greece', 'Hungary', 'Iceland', 'Ireland', 'Italy', 'Latvia', 'Liechtenstein', 'Lithuania', 'Luxembourg', 'Malta', 'Monaco', 'Netherlands', 'Norway', 'Poland', 'Portugal', 'Romania', 'San Marino', 'Serbia', 'Slovenia', 'Spain', 'Sweden', 'Switzerland', 'United Kingdom']
```

```python
a = sorted([country for country, info in countries.items() if info['Continent'] == 'Europe'])
print(a)
```

Output:
```
['Albania', 'Andorra', 'Austria', 'Belgium', 'Bulgaria', 'Croatia', 'Czech Republic', 'Denmark', 'Estonia', 'Finland', 'France', 'Germany', 'Greece', 'Hungary', 'Iceland', 'Ireland', 'Italy', 'Latvia', 'Liechtenstein', 'Lithuania', 'Luxembourg', 'Malta', 'Monaco', 'Netherlands', 'Norway', 'Poland', 'Portugal', 'Romania', 'San Marino', 'Serbia', 'Slovenia', 'Spain', 'Sweden', 'Switzerland', 'United Kingdom']
```

```python
'''def get_europe_countries(countries):
    list_europe = []
    list_europe = sorted([country for country, info in countries.items() if info['Continent'] == 'Europe'])
    return list_europ'''
```

Output:
```
"def get_europe_countries(countries):\n    list_europe = []\n    list_europe = sorted([country for country, info in countries.items() if info['Continent'] == 'Europe'])\n    return list_europ"
```

<!-- grade_id: cell-48bf94a05b056aa0, points: 2, grade: true, solution: false -->
```python
countries_list = get_europe_countries(countries)
assert isinstance(countries_list, list), "Not correct. Keep trying."
assert len(countries_list) == 35, "Not correct. Keep trying."
assert countries_list[0] == "Albania", "Not correct. Keep trying."
assert countries_list[5] == "Croatia", "Not correct. Keep trying."
assert countries_list[34] == "United Kingdom", "Not correct. Keep trying."
source = inspect.getsource(get_europe_countries)
assert "for" in source, "Not correct. Keep trying."
assert "[" in source, "Not correct. Keep trying."
assert "in" in source, "Not correct. Keep trying."
assert ".sort(" in source, "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries = countries_backup
```

Output:
```
---- All asserts passed ----
```

## Question 5

Complete the function `get_literacy_levels_by_continent`. This function should:
- Receive the `countries` dictionary as argument;
- Receive a continent name as argument;
- For the countries in the continent (received as argument), compute the literacy level as:
 - literacy in [0, 25[ % - VERY_LOW;
 - literacy in [25, 50[ % - LOW;
 - literacy in [50, 70[ % - MEDIUM;
 - literacy in [70, 90[ % - HIGH;
 - literacy in [90, 100] % - VERY_HIGH;
- Return a list of tuples like: [(country_1, literacy_1, literacy_level_1), ..., (country_n, literacy_n, literacy_level_n), ...];

```python
countries
```

Output:
```
{'Afghanistan': {'Continent': 'Asia',
  'Population': 31056997.0,
  'Area': 1677018.5,
  'GDP': 700.0,
  'Literacy': 36.0,
  'Phones': 3.22},
 'Albania': {'Continent': 'Europe',
  'Population': 3581655.0,
  'Area': 74457.0,
  'GDP': 4500.0,
  'Literacy': 86.5,
  'Phones': 71.2},
 'Algeria': {'Continent': 'Africa',
  'Population': 32930091.0,
  'Area': 6168682.8,
  'GDP': 6000.0,
  'Literacy': 70.0,
  'Phones': 78.1},
 'American Samoa': {'Continent': 'Oceania',
  'Population': 57794.0,
  'Area': 515.4,
  'GDP': 8000.0,
  'Literacy': 97.0,
  'Phones': 259.54},
 'Andorra': {'Continent': 'Europe',
  'Population': 71201.0,
  'Area': 1212.1,
  'GDP': 19000.0,
  'Literacy': 100.0,
  'Phones': 497.18},
 'Angola': {'Continent': 'Africa',
  'Population': 12127071.0,
  'Area': 3228940.5,
  'GDP': 1900.0,
  'Literacy': 42.0,
  'Phones': 7.78},
 'Anguilla': {'Continent': 'America',
  'Population': 13477.0,
  'Area': 264.2,
  'GDP': 8600.0,
  'Literacy': 95.0,
  'Phones': 460.04},
 'Antigua & Barbud
```

```python
'''
list_cont = []
    if info['Continent'] == 'Europe':
        if info['Literacy'] >= 0 and info['Literacy'] < 25:
            literacy_level = 'VERY_LOW'
        elif info['Literacy'] >= 25 and info['Literacy'] < 50:
            literacy_level = 'LOW'
        elif info['Literacy'] >= 50 and info['Literacy'] < 70:
            literacy_level = 'MEDIUM'
        elif info['Literacy'] >= 70 and info['Literacy'] < 90:
            literacy_level = 'HIGH'
        else:
            literacy_level = 'VERY_HIGH'
            
            
        list_cont.append((country, info['Literacy'], literacy_level))
        
print(list_cont)
'''
```

Output:
```
"\nlist_cont = []\n    if info['Continent'] == 'Europe':\n        if info['Literacy'] >= 0 and info['Literacy'] < 25:\n            literacy_level = 'VERY_LOW'\n        elif info['Literacy'] >= 25 and info['Literacy'] < 50:\n            literacy_level = 'LOW'\n        elif info['Literacy'] >= 50 and info['Literacy'] < 70:\n            literacy_level = 'MEDIUM'\n        elif info['Literacy'] >= 70 and info['Literacy'] < 90:\n            literacy_level = 'HIGH'\n        else:\n            literacy_level = 'VERY_HIGH'\n            \n            \n        list_cont.append((country, info['Literacy'], literacy_level))\n        \nprint(list_cont)\n"
```

<!-- grade_id: cell-3911f29ed53cd968, points: 0, grade: false, solution: true -->
```python
def get_literacy_levels_by_continent(countries, continent):
    list_cont = []
    literacy_level = 0
    for country, info in countries.items():
        if info['Continent'] == continent:
            
            if info['Literacy'] >= 0 and info['Literacy'] < 25:
                literacy_level = 'VERY_LOW'
            elif info['Literacy'] >= 25 and info['Literacy'] < 50:
                literacy_level = 'LOW'
            elif info['Literacy'] >= 50 and info['Literacy'] < 70:
                literacy_level = 'MEDIUM'
            elif info['Literacy'] >= 70 and info['Literacy'] < 90:
                literacy_level = 'HIGH'
            else:
                literacy_level = 'VERY_HIGH'
            list_cont.append((country, info['Literacy'], literacy_level))
    return list_cont
```

```python
x = get_literacy_levels_by_continent(countries, 'Europe')
print(x)
```

Output:
```
[('Albania', 86.5, 'HIGH'), ('Andorra', 100.0, 'VERY_HIGH'), ('Austria', 98.0, 'VERY_HIGH'), ('Belgium', 98.0, 'VERY_HIGH'), ('Bulgaria', 98.6, 'VERY_HIGH'), ('Croatia', 98.5, 'VERY_HIGH'), ('Czech Republic', 99.9, 'VERY_HIGH'), ('Denmark', 100.0, 'VERY_HIGH'), ('Estonia', 99.8, 'VERY_HIGH'), ('Finland', 100.0, 'VERY_HIGH'), ('France', 99.0, 'VERY_HIGH'), ('Germany', 99.0, 'VERY_HIGH'), ('Greece', 97.5, 'VERY_HIGH'), ('Hungary', 99.4, 'VERY_HIGH'), ('Iceland', 99.9, 'VERY_HIGH'), ('Ireland', 98.0, 'VERY_HIGH'), ('Italy', 98.6, 'VERY_HIGH'), ('Latvia', 99.8, 'VERY_HIGH'), ('Liechtenstein', 100.0, 'VERY_HIGH'), ('Lithuania', 99.6, 'VERY_HIGH'), ('Luxembourg', 100.0, 'VERY_HIGH'), ('Malta', 92.8, 'VERY_HIGH'), ('Monaco', 99.0, 'VERY_HIGH'), ('Netherlands', 99.0, 'VERY_HIGH'), ('Norway', 100.0, 'VERY_HIGH'), ('Poland', 99.8, 'VERY_HIGH'), ('Portugal', 93.3, 'VERY_HIGH'), ('Romania', 98.4, 'VERY_HIGH'), ('San Marino', 96.0, 'VERY_HIGH'), ('Serbia', 93.0, 'VERY_HIGH'), ('Slovenia', 99.7, 'VE
```

<!-- grade_id: cell-fa41a51f9ff00d7c, points: 2, grade: true, solution: false -->
```python
for continent in ["Europe","Africa"]:
    literacy_tuple = get_literacy_levels_by_continent(countries, continent)    
    for literacy in literacy_tuple:
        assert len(literacy) == 3, "Not correct. Keep trying."
        assert isinstance(literacy[0], str), "Not correct. Keep trying."
        assert isinstance(literacy[1], float), "Not correct. Keep trying."
        assert isinstance(literacy[2], str), "Not correct. Keep trying."

    if continent == "Europe":
        assert len(literacy_tuple) == 35, "Not correct. Keep trying."
        assert literacy_tuple[21][0] == 'Malta', "Not correct. Keep trying."
        assert math.isclose(literacy_tuple[21][1], 92.8, abs_tol=0.1), \
               "Not correct. Keep trying."
        assert literacy_tuple[21][2] == 'VERY_HIGH', "Not correct. Keep trying."
        assert literacy_tuple[1][2] == 'VERY_HIGH', "Not correct. Keep trying."
        assert literacy_tuple[0][2] == 'HIGH', "Not correct. Keep trying."
        assert literacy_tuple[-1][2] == 'VERY_HIGH', "Not correct. Keep trying."

    if continent == "Africa":
        assert len(literacy_tuple) == 55, "Not correct. Keep trying."
        assert literacy_tuple[14][0] == 'Djibouti', "Not correct. Keep trying."
        assert math.isclose(literacy_tuple[14][1], 67.9, abs_tol=0.1), \
               "Not correct. Keep trying."
        assert literacy_tuple[14][2] == 'MEDIUM', "Not correct. Keep trying."
        assert literacy_tuple[1][2] == 'LOW', "Not correct. Keep trying."
        assert literacy_tuple[0][2] == 'HIGH', "Not correct. Keep trying."
        assert literacy_tuple[-1][2] == 'VERY_HIGH', "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries = countries_backup
```

Output:
```
---- All asserts passed ----
```

## Question 6

Complete the function `get_country_codes`. This function should:
- Receive the `countries` dictionary as argument;
- Build a list with all the countries names;
- Using `map` and a `lambda` function, convert the country names in country codes by selecting the 3 first letters of the country;
- Returns this list;

Note: Don't worry about duplicates on the result.

```python
#con =list(countries.keys())
#con =  map(lambda x: x[0:3], con)
#con = list(con)
```

```python
'''
con2 = map(lambda x: x[0:3], list(countries.keys()))
con2 = list(con2)
'''
```

Output:
```
'\ncon2 = map(lambda x: x[0:3], list(countries.keys()))\ncon2 = list(con2)\n'
```

```python
#con3 = = list(map(lambda x: x[0:3], list(countries.keys())))
```

<!-- grade_id: cell-bd068242631b4909, points: 0, grade: false, solution: true -->
```python
def get_country_codes(countries):
    con = list(map(lambda x: x[0:3], list(countries.keys())))
    return con
```

<!-- grade_id: cell-5f1eaf642616c8a1, points: 2, grade: true, solution: false -->
```python
country_codes = get_country_codes(countries) 
assert len(country_codes) == 194, "Not correct. Keep trying."
for code in country_codes:
    assert isinstance(code, str), "Not correct. Keep trying."
    assert len(code) == 3, "Not correct. Keep trying."
assert country_codes[134] == "Pak", "Not correct. Keep trying."
assert country_codes[62] == "Fra", "Not correct. Keep trying."

source = inspect.getsource(get_country_codes)
assert "map" in source, "Not correct. Keep trying."
assert "lambda" in source, "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries = countries_backup
```

Output:
```
---- All asserts passed ----
```

## Question 7

The `Country` class is going to be used to represent countries and to return information about them.

### Question 7.1

Complete the `get_population_in_millions` method. This method should:
- Return the population of the country that the class represents;
- The population should be returned in millions with 2 decimal digits;

<!-- grade_id: cell-9b9ae4ba8e61e508, points: 0, grade: false, solution: true -->
```python
class Country:
    def __init__(self, country_name, population, continent):
        self.country_name = country_name
        self.population = population
        self.continent = continent

    def get_population_in_millions(self):
        return round((self.population*1e-6),2)
```

```python
country1 = Country("Penguinea", 12345678, "Antarctica")
```

```python
country1.get_population_in_millions()
```

Output:
```
12.35
```

<!-- grade_id: cell-9e28e403f317f23d, points: 2, grade: true, solution: false -->
```python
country = Country("Penguinea", 12345678, "Antarctica")
assert isinstance(country, Country), "Not correct. Keep trying."
assert isinstance(country.get_population_in_millions(), float), "Not correct. Keep trying."
assert math.isclose(country.get_population_in_millions(), 12.35, abs_tol=0.01),"Not correct. Keep trying."

country = Country("Walrussia", 8344567, "Antarctica") 
assert math.isclose(country.get_population_in_millions(), 8.34, abs_tol=0.01), "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries = countries_backup
```

Output:
```
---- All asserts passed ----
```

## Question 7.2

Complete the `get_country_population`. This function should:
- Receive the `countries` dictionary as argument;
- Receive a country name as argument;
- Try to find the country (received as argument) in the countries dictionary;
- If the country doesn’t exist in the countries dictionary, catch the KeyError, print a statement with the information that there is no information for that country and return `None`;
- If the country exists, create an object of the `Country` class;
- Call the `get_population_in_millions` method for the country;
- Return the population in millions value;

```python
countries['Turkey']
```

Output:
```
{'Continent': 'Asia',
 'Population': 70413958.0,
 'Area': 2021694.4,
 'GDP': 6700.0,
 'Literacy': 86.5,
 'Phones': 269.52}
```

<!-- grade_id: cell-803b263dbee2e300, points: 0, grade: false, solution: true -->
```python
def get_country_population(countries, country_name):
    try:
        if country_name in countries:
            country_name = Country(country_name = country_name,
                    population = countries[country_name]['Population'],
                    continent = countries[country_name]['Continent'])
            return country_name.get_population_in_millions()
            
    except KeyError as countries:
        print(f'Country {country_name} Not Found in countries Dictionary')
        return None
        
        
```

<!-- grade_id: cell-0d0db73f5256c8e5, points: 6, grade: true, solution: false -->
```python
assert math.isclose(get_country_population(countries, "Turkey"), 70.41, abs_tol=0.01), "Not correct. Keep trying."
assert get_country_population(countries, "Narnia") is None,  "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries = countries_backup
```

Output:
```
---- All asserts passed ----
```

### Submit your work!

To submit your work, [follow the steps here, in the step \"Grading the Exercise Notebook\"!](https://github.com/LDSSA/ds-prep-course-2023/blob/main/weekly-workflow.md#link-to-grading)